---
title       : Wrap up & next steps
description : Setup your own data.world account and apply what you've learned and more!

--- type:NormalExercise lang:python xp:100 skills:2 key:d55f0708ee
## Wrap up & next steps

Congrats! You now know multiple ways of pulling data from data.world into your local environment using our Python SDK. But don't stop there! The Python SDK has additional functionality for uploading and modifying data for easy storage as well as sharing and collaborating with other members. 

data.world is completely free, so signup at [https://data.world/](https://data.world/) and start applying your new skills. We've also put together a handy jupyter notebook with all of the things we've covered in this tutorial, as well as some of the more advanced functionality. 

Download the notebook [here](https://query.data.world/s/59dlfssv3zwf1k2l8k18gg9z3), and follow the below steps for installing the data.world Python SDK on your machine:

1 ) Install data.world's Python package, including optional `pandas` dependencies:

`pip install datadotworld[pandas]`

2 ) Run the dw command-line tool to configure your machine:

`dw configure`

3 ) Copy/Paste in your data.world API token when prompted. Find your token at [https://data.world/settings/advanced](https://data.world/settings/advanced)


*** =instructions
- Setup your data.world account at `https://data.world/`.
- Install the datadotworld Python package to your local machine using `pip install datadotworld[pandas]`.
- Download the [jupyter notebook](https://query.data.world/s/59dlfssv3zwf1k2l8k18gg9z3) to quickly apply the things you've learned in this tutorial.
- Use data.world to find more interesting data, store and share your own data and projects, and collaborate with other community members!

*** =hint

*** =pre_exercise_code
```{python}

```

*** =sample_code
```{python}

```

*** =solution
```{python}

```

*** =sct
```{python}

```